---
title: Test
description: Test
---

## Insert exercise title here

```yaml
type: NormalExercise
key: d42f6d0ac5
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 7dcadcaf05
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 60b233bb9d
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: TabExercise
key: e9b486b96a
xp: 100
```



`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


---

## Insert exercise title here

```yaml
type: BulletExercise
key: 4224f5abb3
xp: 100
```



`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```
